## Project description

Goal: Create a software to compare structural similarity of two Python codes. Input file extension is ``.txt``. There are two files, ``file1.txt`` and ``file2.txt``, available here as a sample input.
The first line of the input file is student information. The actual Python code starts at line 2. 

1\. Define a function, ``get_student_info()``, that returns student's information. **(2 point)**<br>
> Sample run:  ``get_student_info('file1.txt')``: 

```Python
Out[1]: 'student 1'```
<br>
2\. Define ``read_code()`` to read Python code in the file. This function returns the body of the code(without student information). **(2 point)**<br>
> Sample run:  ``read_code('file1.txt')``:

```Python
Out[2]:
num_sum = 0
for x in range(10):
   num_sum += x

print(num_sum)```
<br>
3\. Define a function ``python_tokenize()``, to split input codes into a list of words. Punctuations should be considered as a word in the output. This function takes text as an input, which is an output of read_code() function. **(3 point)**<br>
> Sample run: ``python_tokenize(text1)``:

```Python
Out[3]: ['num_sum', '=', '0', 'for', 'x', 'in', 'range', '(', '10', ')', ':', 'num_sum', '+=', 'x', 'print', '(', 'num_sum', ')']```
<br>
4\. Define a function ``filter_stopwords()`` to remove stop words from the tokenized list. This function takes tokenized list as an input, which is an output of python_tokenize() function. We have also provided list of stopwords to be removed. **(2 point)**<br>
> Sample run: ``filter_stopwords(tokens)``:

```Python
Out[4]: ['num_sum', '=', '0', 'x', 'range', '(', '10', ')', ':', 'num_sum', '+=', 'x', 'print', '(', 'num_sum', ')']```
<br>
5\. Define a function ``function_list()`` to return the list all the function names in the code. This function takes text as an input, which is an output of read_code() function. **(3 point)**<br>
> Sample run: ``function_list(text1)``: 

```Python
Out[5]: ['range', 'print']```
<br>
6\. Define a function ``variable_list()`` to return the list of all variable names in the code. This function takes text as an input, which is an output of read_code() function. **(3 point)**<br>
> Sample run: ``variable_list(text1)``:

```Python
Out[6]: ['num_sum']```
<br>
7\. Define a function ``cosine_similarity()`` to return the cosine similarity of two codes. This function takes two tokenized lists as input, which are outputs of python_tokenize() function. **(5 point)**

You need to compute **cosine similarity** of two list of tokens. Read [here] to learn about cosine similarity. 

As definition of cosine similarity uses norm and dot product of the vectors, you need to first define these two functions. Then you can call them while defining ``cosine_similarity()``. Here is some hint:
    
   ```Python
    make a word_list vector
    vector1 = frequenct vector of sentence 1
    vector2 = frequency vector of sentence 2
    cosine_similarty = dot_product(vector1. vector2)/norm(vector1).norm(vector2)
   ```
> Sample run: ``cosine_similarity(python_tokenize(text1), python_tokenize(text2))``:

```Python
Out[7]: 0.598764159346999```
<br>
8\. Define a function ``report()``. This function accepts two text files as input arguments and creates a report file called ``similarity_report.txt``. **(5 point)**

Funtion should perform following tasks in this order:
1. Create new file called ``similarity_report.txt`` in the working directory or overwrite the file if file with same name already exists.
2. Write report sentences to this file.Check sample file for required format of report.
3. Close the file.
4. This function won't return anything such as report text. It will just write report to the file.

> Sample run: ``report('file1.txt', 'file2.txt')``:<br><br> No output below cell. When you call the function using this code  , ``similarity_report.txt`` file that it creates should match **EXACTLY** with ``similarity_report_sample.txt`` file given to you. Otherwise your function won't pass the test case given in the cell below.

Hint - This function will use all seven funtions above.
<br>
<br>
**Important notes:** 
- You can call previously defined functions while defining the current function, but you cannot define any new fuctions that are not given above.
- Don't change function names or their parameters.

[here]:https://en.wikipedia.org/wiki/Cosine_similarity

<!-- grade_id: Funtion1, points: 0, grade: false, solution: true -->
```python
def get_student_info(file_name):
    # YOUR CODE HERE
    f = open(file_name)
    contents = f.readline()
    f.close()
    return contents[:-1]

# get_student_info('file1.txt')
```

<!-- grade_id: Test1, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal
result1 = 'student 1'
assert_equal(get_student_info('file1.txt'),result1)

```

<!-- grade_id: Function2, points: 0, grade: false, solution: true -->
```python
def read_code(file_name):
    # YOUR CODE HERE
    f = open(file_name)
    contents = f.read()
    f.close()
    return contents[len(get_student_info(file_name))+1:]

# read_code('file1.txt')
```

<!-- grade_id: Test2, points: 2, grade: true, solution: false -->
```python
result2 = '''num_sum = 0
for x in range(10):
   num_sum += x

print(num_sum)'''
assert_equal(read_code('file1.txt'),result2)

```

<!-- grade_id: Function3, points: 0, grade: false, solution: true -->
```python
def python_tokenize(text):
    # YOUR CODE HERE
    lst1 = text.split()
    lst2 = []
    for i in lst1:
        if i.find('(') != -1:
            a = i.find('(')
            b = i.find(')')
            test = i[:a]+' '+i[a]+' '+i[a+1:b]+' '+i[b]+' '+i[b+1:]
            lst2 = lst2 + test.split()
        elif i.find('<') != -1 or i.find(':') != -1:
            a = i.find('<')
            b = i.find(':')
            test = i[:a]+' '+i[a]+' '+i[a+1:b]+' '+i[b]+' '+i[b+1:]
            lst2 = lst2 + test.split()
        elif i.find('=1') != -1:
            test = i.replace('=1','= 1')
            lst2 = lst2 + test.split()
        else:
            lst2.append(i)
    return lst2

# text = read_code('file2.txt')
# python_tokenize(text)
```

<!-- grade_id: Test3, points: 3, grade: true, solution: false -->
```python
text1 = read_code('file1.txt')
result3 = ['num_sum', '=', '0', 'for', 'x', 'in', 'range', '(', '10', ')', ':', 'num_sum', '+=', 'x', 'print', '(', 'num_sum', ')']
assert_equal(python_tokenize(text1),result3)

text2 = read_code('file2.txt')
result3_2 = ['num', '=', '0', 'num_sum', '=', '0', 'while', 'num', '<', '10', ':', 'num_sum', '+=', 'num', 'num', '+=', '1', 'print', '(', 'num_sum', ')']

assert_equal(python_tokenize(text2), result3_2)
```

<!-- grade_id: Function4, points: 0, grade: false, solution: true -->
```python
python_stopping_word = ['False', 'class', 'finally', 'is', 'return', 'None', 'continue', 'for', 'lambda', 'try', 'True', 
                        'def','from', 'nonlocal', 'while','and', 'del', 'global', 'not', 'with', 'as', 'elif', 'if', 'or', 
                        'yield', 'assert', 'else', 'import', 'pass', 'break', 'except', 'in', 'raise']
def filter_stopwords(tokens):
    # YOUR CODE HERE
    for i in python_stopping_word:
        if i in tokens:
            tokens.remove(i)
    return tokens
```

<!-- grade_id: Test4, points: 2, grade: true, solution: false -->
```python
tokens = python_tokenize(text1)
result4 = ['num_sum', '=', '0', 'x', 'range', '(', '10', ')', ':', 'num_sum', '+=', 'x', 'print', '(', 'num_sum', ')']
assert_equal(filter_stopwords(tokens),result4)

```

<!-- grade_id: Function5, points: 0, grade: false, solution: true -->
```python
def function_list(text):
    # YOUR CODE HERE
    lst1 = text.split()
    lst2 = []
    for i in lst1:
        if i.find('(') and i.find(')') != -1:
            p1 = i.find('(')
            p2 = i.find(')')
            lst2.append(i[:p1])
    return lst2

# function_list(read_code('file1.txt'))
```

<!-- grade_id: Test5, points: 3, grade: true, solution: false -->
```python
result5 = ['range', 'print']
assert_equal(function_list(text1),result5)

```

<!-- grade_id: Function6, points: 0, grade: false, solution: true -->
```python
def variable_list(text):
    # YOUR CODE HERE
    assignment_list = ['=', '+=', '-=', '/=', '*=']
    j = 1
    lst1 = text.split()
    lst2 = []
    for i in lst1:
        if j == len(lst1):
            break
        if lst1[j] in assignment_list:
            lst2.append(i)
        j=j+1
    return list(set(lst2))
```

<!-- grade_id: Test6, points: 3, grade: true, solution: false -->
```python
result6 = ['num_sum']
assert_equal(variable_list(text1),result6)

```

<!-- grade_id: Function7, points: 0, grade: false, solution: true -->
```python
def norm(vec):
    # YOUR CODE HERE
    import math
    s = 0
    for i in vec:
        s += i**2
    return math.sqrt(s)
    
def dot(vec1, vec2):
    # YOUR CODE HERE
    s = 0
    for i in range(len(vec1)):
        s += vec1[i]*vec2[i]
    return s
    
def cosine_similarity(list1, list2):
    # YOUR CODE HERE
    vector1 = []
    vector2 = []
    word_list = list(set(list1+list2))
    for i in range(len(word_list)):
        vector1.append(list1.count(word_list[i]))
        vector2.append(list2.count(word_list[i]))
    cosine_similarity = dot(vector1, vector2)/(norm(vector1)*norm(vector2))
    return cosine_similarity
```

<!-- grade_id: Test7, points: 5, grade: true, solution: false -->
```python
result7 = 0.598764159346999
sim1 = cosine_similarity(python_tokenize(text1), python_tokenize(text2))
assert_equal(sim1, result7)

result7_2 = 0.6267831705280087
sim2 = cosine_similarity(filter_stopwords(python_tokenize(text1)), filter_stopwords(python_tokenize(text2)))
assert_equal(sim2, result7_2)
```

<!-- grade_id: Function8, points: 0, grade: false, solution: true -->
```python
def report(file1, file2):
    # YOUR CODE HERE
    f = open('similarity_report.txt', 'w')
    f.write(get_student_info(file1)+'\n')
    f.write('variables: '+ str(variable_list(read_code(file1)))+'\n')
    f.write('functions: '+ str(function_list(read_code(file1)))+'\n')
    f.write(get_student_info(file2)+'\n')
    f.write('variables: '+ str(variable_list(read_code(file2)))+'\n')
    f.write('functions: '+ str(function_list(read_code(file2)))+'\n')
    f.write('Overall similarity: '+ str(cosine_similarity(python_tokenize(read_code(file1)), python_tokenize(read_code(file2))))+'\n')
    f.write('Similarity without stopping words: '+ str(cosine_similarity(filter_stopwords(python_tokenize(read_code(file1))), filter_stopwords(python_tokenize(read_code(file2)))))+'\n')
    f.close()
```

<!-- grade_id: Test8, points: 5, grade: true, solution: false -->
```python
report('file1.txt', 'file2.txt')
outputfile = open('similarity_report.txt', 'r')
outputreport = outputfile.read()
samplefile = open('similarity_report_sample.txt', 'r')
samplereport = samplefile.read()

assert_equal(outputreport, samplereport)
```

Output:
```
c:\python37\lib\unittest\case.py in assertMultiLineEqual(self, first, second, msg)
   1218             diff = '\n' + ''.join(difflib.ndiff(firstlines, secondlines))
   1219             standardMsg = self._truncateMessage(standardMsg, diff)
-> 1220             self.fail(self._formatMessage(msg, standardMsg))
   1221 
   1222     def assertLess(self, a, b, msg=None):

c:\python37\lib\unittest\case.py in fail(self, msg)
    678     def fail(self, msg=None):
    679         """Fail immediately, with the given message."""
--> 680         raise self.failureException(msg)
    681 
    682     def assertFalse(self, expr, msg=None):

AssertionError: "stud[84 chars]['num_sum', 'num']\nfunctions: ['print']\nOver[86 chars]87\n" != "stud[84 chars]['num', 'num_sum']\nfunctions: ['print']\nOver[86 chars]87\n"
  student 1
  variables: ['num_sum']
  functions: ['range', 'print']
  student 2
- variables: ['num_sum', 'num']
?                      -------
+ variables: ['num', 'num_sum']
?             ++++
```